<problem display_name="samplePRO">
<script type="loncapa/python">
<![CDATA[
import json
from decimal import Decimal
def vglcfn(e, ans):
    '''
    par is a dictionary containing two keys, "answer" and "state"
    The value of answer is the JSON string returned by getGrade
    The value of state is the JSON string returned by getState
    '''
    par = json.loads(ans)
    answer = json.loads(par["answer"])
    total = answer["total"]
    correct = answer["correct"]
    undo = answer["undo"]
    fix = answer["fix"]
    student = answer["student"]
    grade_decimal = (correct - undo - fix)  / total
    return {'msg':'Good', 'ok': True, 'grade_decimal': grade_decimal}
]]>
</script>
<customresponse cfn="vglcfn">
  <jsinput gradefn="samplePRO.getGrade"
    get_statefn="samplePRO.getState"
    set_statefn="samplePRO.setState"
    width="600"
    height="300"
    html_file="https://opendsax.herokuapp.com/AV/Sample/samplePRO.html"
    sop="false"/>
</customresponse>
</problem>