<problem display_name="quicksortPRO">
<script type="loncapa/python">
<![CDATA[
import json
def vglcfn(e, ans):
    '''
    par is a dictionary containing two keys, "answer" and "state"
    The value of answer is the JSON string returned by getGrade
    The value of state is the JSON string returned by getState
    '''
    par = json.loads(ans)
    answer = json.loads(par["answer"])
    total = answer["total"]
    correct = answer["correct"]
    undo = answer["undo"]
    fix = answer["fix"]
    student = answer["student"]
    score = correct - undo - fix
    if score < 0:
       score = 0
    grade_decimal = score / total
    return {'msg':'Good', 'ok': True, 'grade_decimal': grade_decimal}
]]>
</script>
  <p>
Here is a complete proficiency exercise to see how well you understand Quicksort.
</p>
  <customresponse cfn="vglcfn">
    <jsinput gradefn="quicksortPRO.getGrade"
      get_statefn="quicksortPRO.getState"
      set_statefn="quicksortPRO.setState"
      width="792" 
      height="512" 
      html_file="https://opendsax.herokuapp.com/AV/Sorting/quicksortPRO.html" 
      sop="false"/>
  </customresponse>
</problem>