<problem display_name="webGLDemo">
<script type="loncapa/python">
<![CDATA[
import json
def vglcfn(e, ans):
  '''
  par is a dictionary containing two keys, "answer" and "state"
  The value of answer is the JSON string returned by getGrade
  The value of state is the JSON string returned by getState
  '''
  par = json.loads(ans)
  # We can use either the value of the answer key to grade
  answer = json.loads(par["answer"])
  return answer["cylinder"] and not answer["cube"]
  '''
  # Or we could use the value of the state key
  state = json.loads(par["state"])
  selectedObjects = state["selectedObjects"]
  return selectedObjects["cylinder"] and not selectedObjects["cube"]
  '''
]]>
</script>
<p>
Here is a complete proficiency exercise to see how well you understand Quicksort.
</p>
<customresponse cfn="vglcfn">
  <jsinput gradefn="WebGLDemo.getGrade"
    get_statefn="WebGLDemo.getState"
    set_statefn="WebGLDemo.setState"
    initial_state='{"selectedObjects":{"cube":true,"cylinder":false}}'
    width="792"
    height="512"
    html_file="http://opendsax.local/av/sorting/quicksortPRO.html"
    sop="true"/>
</customresponse>
</problem>